Before you turn this lab in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
DRIVER = ""
NAVIGATOR = ""
```

# Optical Character Recognition Lab

Welcome to the optical character recognition lab! By the end of this lab you will have

- *Trained a neural network to perform optical character recognition*

Let's get started!

```python
def passed(): print('✅')
```

# Load Data

The following code loads and normalizes the MNIST dataset and filters it down to a set containing only exemplar images.

```python
from IPython.display import display
from keras.preprocessing.image import array_to_img
from keras.datasets import mnist
import numpy as np
np.random.seed(1337)

[X_mnist, y_mnist], _ = mnist.load_data()
X_mnist = X_mnist/255.
digit_idxs = [np.argwhere(y_mnist == d).flatten()[0] for d in range(10)]
X_mnist = X_mnist[digit_idxs]

for x in X_mnist:
    img = array_to_img(np.expand_dims(x, axis=-1))
    display(img)
```

# Generate Training Data

The following code generates sequences of the exemplar MNIST digits along with their labels and provides the option to change two hyperparameters.

```python
NB_TRAIN = 90
T = 2
```

```python
import itertools
from keras.preprocessing.image import array_to_img

H = W = 28
X, y = np.zeros([NB_TRAIN, H, W*T]), [None]*NB_TRAIN
i = 0
D = set()
while i < NB_TRAIN:
    digits = np.random.choice(10, size=T, replace=False)
    if tuple(digits) in D:
        continue
    D.add(tuple(digits))
    X[i], y[i] = np.hstack([x for x in X_mnist[digits]]), digits
    i += 1 
X = np.expand_dims(X, axis=-1)

Y = np.zeros([NB_TRAIN, T, 10])
for i, ys in enumerate(y):
    for j, y_ in enumerate(ys):
        Y[i, j, y_] = 1

for idx in range(5):
    x = X[idx]
    img_ = array_to_img(X[idx])
    print(y[idx])
    display(img_)
    
X_train, Y_train = X[NB_TRAIN//10:], Y[NB_TRAIN//10:]
X_val, Y_val = X[:NB_TRAIN//10], Y[:NB_TRAIN//10]

print(X_train.shape, Y_train.shape)
print(X_val.shape, Y_val.shape)
```

# Task

- Define and compile a deep learning model to perform optical character recognition on the MNIST sequence dataset

# Requirements

- Save your model into a python variable called `model`

<!-- grade_id: task_solution, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Model Tests

<!-- grade_id: task_tests, points: 1, grade: true, solution: false -->
```python
assert model.input_shape == (None,) + X[0].shape
assert model.output_shape == (None,) + Y[0].shape
assert model.count_params() > 100_000
assert model.loss == 'categorical_crossentropy'
assert model.metrics == ['accuracy']
assert model.predict(X[:1])[0].shape == Y[0].shape

passed()
```

# Task

- Optimize `model` on `X_train` and `Y_train`

<!-- grade_id: fit_solution, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Performance Tests

<!-- grade_id: test_gt_10, points: 1, grade: true, solution: false -->
```python
loss, acc = model.evaluate(X_val, Y_val)
assert acc > 0.10
passed()
```

<!-- grade_id: test_gt_15, points: 1, grade: true, solution: false -->
```python
loss, acc = model.evaluate(X_val, Y_val)
assert acc > 0.15
passed()
```

<!-- grade_id: test_gt_20, points: 1, grade: true, solution: false -->
```python
loss, acc = model.evaluate(X_val, Y_val)
assert acc > 0.20
passed()
```

<!-- grade_id: test_gt_25, points: 1, grade: true, solution: false -->
```python
loss, acc = model.evaluate(X_val, Y_val)
assert acc > 0.25
passed()
```

<!-- grade_id: test_gt_30, points: 1, grade: true, solution: false -->
```python
loss, acc = model.evaluate(X_val, Y_val)
assert acc > 0.30
passed()
```

<!-- grade_id: test_gt_50, points: 1, grade: true, solution: false -->
```python
loss, acc = model.evaluate(X_val, Y_val)
assert acc > 0.50
passed()
```

<!-- grade_id: test_gt_75, points: 1, grade: true, solution: false -->
```python
loss, acc = model.evaluate(X_val, Y_val)
assert acc > 0.75
passed()
```

<!-- grade_id: test_gt_90, points: 1, grade: true, solution: false -->
```python
loss, acc = model.evaluate(X_val, Y_val)
assert acc > 0.90
passed()
```