Before you turn this lab in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
DRIVER = ""
NAVIGATOR = ""
```

# TensorFlow + Keras Lab

Welcome to the TensorFlow + Keras lab! By the end of this lab you will have

- Optimized and evaluated a keras model with TensorFlow

Let's get started!

# Autograder Variables

The following code defines variables owned by the autograder.

```python
def passed(): print('✅')
```

# Set the Stage

The following code loads data and normalizes it and defines a TensorFlow session and registers it with keras and defines a keras model.

```python
import tensorflow as tf
from keras.datasets import boston_housing
from keras.models import Sequential
from keras import backend as K
from keras.layers import Dense

[X_boston, y_boston], _ = boston_housing.load_data()
X_boston = (X_boston - X_boston.mean(axis=0)) / X_boston.std(axis=0)
y_boston = (y_boston - y_boston.mean()) / y_boston.std()

sess = tf.Session()
K.set_session(sess)

model = Sequential()
model.add(Dense(units=1, input_shape=X_boston[0].shape))
model.summary()
```

# Task

- Define tensors `labels` and `loss` which hold your labels and aggregated loss for each training example in your minibatch respectively

<!-- grade_id: tensors_solution, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Tensor Tests

<!-- grade_id: tensors_tests, points: 1, grade: true, solution: false -->
```python
assert loss.shape.as_list() == []
assert type(loss) == tf.Tensor
assert type(labels) == tf.Tensor
assert labels.op.type == 'Placeholder'
assert loss.dtype == tf.float32
assert loss.shape == []
assert loss.op.type == 'Mean'

passed()
```

# Task

- Define a TensorFlow optimizer and initialize all variables

# Requirments

- Define a tensor `train_step` which runs a step of `Adam` to minimize `loss`
- Define a tensor `init` which initializes all `tf.Variable`s and run it

<!-- grade_id: optimizer_solution, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Optimizer Tests

<!-- grade_id: optimizer_tests, points: 1, grade: true, solution: false -->
```python
assert type(train_step) == tf.Operation
assert type(train_step) == tf.Operation
assert train_step.type == 'NoOp'
assert init.type == 'NoOp'

passed()
```

# Task

- Optimize `model` so the loss gets below 1.0 on `(X_boston, y_boston)`

<!-- grade_id: optimize_solution, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Evaluation Test

<!-- grade_id: evaluation_test, points: 1, grade: true, solution: false -->
```python
best_loss = sess.run(fetches=loss, feed_dict={model.input: X_boston, labels: np.expand_dims(y_boston, axis=-1)})
assert best_loss < 1.0

passed()
```