Before you turn this lab in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
DRIVER = ""
NAVIGATOR = ""
```

# Optimization II Lab

Welcome to the optimization II lab! By the end of this lab you will have

- Performed optimization on a deep learning model with several different optimizers
- Visualized optimizer dynamics via TensorBoard
- Visualized optimizer performance via SacredBoard

Let's get started!

# Unit Test Variables

The following code defines variables that will be used in subsequent unit tests. Do not attempt to redefine any of these variables throughout the notebook!

```python
from IPython.display import HTML

def passed():
    print('✅')
```

## Task

- Define a convolutional neural network trainer `CNNTrainer` in `trainers.py`

## Remarks

- You can either fill in the code below or directly edit `trainers.py`

<!-- grade_id: cnn_solution, points: 0, grade: false, solution: true -->
```python
from trainer import Trainer

class CNNTrainer(Trainer):
    """Convolutional Neural Network Classifier"""

    def build_model(self):
        from keras.models import Sequential
        from keras.layers import Conv2D, Dense, MaxPooling2D, Dropout, Flatten

        model = Sequential()
        
        # YOUR CODE HERE
        raise NotImplementedError()
        
        self.model = model
        
import trainers
trainers.CNNTrainer = CNNTrainer
```

# `CNNTrainer` Tests

<!-- grade_id: cnn_tests, points: 1, grade: true, solution: false -->
```python
import numpy as np
from trainers import CNNTrainer

cnn = CNNTrainer(config={})
cnn.load_data()
cnn.build_model()

nb_conv = len([layer for layer in cnn.model.layers if layer.name.startswith('conv')])
assert nb_conv >= 1
assert cnn.model.input_shape == (None, 28, 28, 1)
assert cnn.model.output_shape == (None, 10)
X_ = np.random.randn(16, 28, 28, 1)
y_pred = cnn.model.predict_classes(X_, verbose=0)
assert np.all(0 <= y_pred) and np.all(y_pred < 10)

passed()
```

## Task

- Optimize your `CNNTrainer` model with every optimizer in keras with the exception of `TFOptimizer`

## Suggestion

- Start with a simple model like `MLRTrainer` and only consider a small number of training examples so you can debug quickly until you're sure everything is working correctly and them experiment with larger slower `CNNTrainer`

## Requirement

- Keep the sacred [Mongo Observer](http://sacred.readthedocs.io/en/latest/observers.html#mongo-observer) so you can view the results afterward in [sacredboard](https://github.com/chovanecm/sacredboard)

<!-- grade_id: runs_solution, points: 0, grade: false, solution: true -->
```python
from train import ex
import keras
import keras.backend as K
from sacred.observers import MongoObserver
mongo_observer = MongoObserver.create()
ex.observers.append(mongo_observer)

# YOUR CODE HERE
raise NotImplementedError()
```

## Task

- Visualize optimizer dynamics with TensorBoard

## Requirements

- Select out plots which track only training loss during your experiment runs in TensorBoard with an appropriate regex (e.g. `loss`)
- Take a screenshot
- Load it into a `IPython.display.Image` object called `tensorboard_screenshot`
- Display it

<!-- grade_id: screenshot_solution, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# TensorBoard Screenshot Test

<!-- grade_id: tensorboard_screenshot_test, points: 1, grade: true, solution: false -->
```python
import IPython

assert type(tensorboard_screenshot) == IPython.core.display.Image
assert 'PNG' in str(tensorboard_screenshot.data) or 'JPG' in str(tensorboard_screenshot.data)

passed()
```

## Task

- Visualize optimizer performance with sacredboard

## Requirements

- Sort the optimizer runs in sacredboard by loss
- Take a screenshot
- Load it into a `IPython.display.Image` object called `sacredboard_screenshot`
- Display it

<!-- grade_id: sacredboard_screenshot_solution, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# SacredBoard Screenshot Test

<!-- grade_id: sacredboard_screenshot_test, points: 1, grade: true, solution: false -->
```python
import IPython

assert type(sacredboard_screenshot) == IPython.core.display.Image
assert 'PNG' in str(sacredboard_screenshot.data) or 'JPG' in str(sacredboard_screenshot.data)

passed()
```

## Question

- Which optimizer did you find worked best? Does that surprise you? If so, why?

<!-- grade_id: best_answer, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

## Question

- What is the intution behind that optimitzer?

<!-- grade_id: intuition_answer, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

# Challenge Activities

- Tune the hyperparameters of the best optimizer
- Do the same thing in TensorFlow